Collaborators: 
     - Zachary Chommala 
     - Bao Nguyen 
     - Alex Ho

# CSCI 3155: Assignment 5

Topics Covered: Operations on inductive definitions, building an interpreter, big step operational semantics and Lettuce.

__Name__: Jonathan Phouminh

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"Tests Passed (1 point)")
    else print(s"Tests Passed ($points points)") 
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1: Translating Lettuce Into Scala (25 points)

In this problem, we will translate Lettuce programs into scala. We will consider the fragment of the
language with Let bindings and if-then-else statements.

$$\newcommand\Expr{\mathbf{Expr}}$$
$$ \begin{array}{rcl}
\Expr &\Rightarrow & \text{Const}(\mathbf{Double}) \\
& | & \text{ConstTrue}\\
&|& \text{ConstFalse}\\
& | & \text{Ident}(\mathbf{String}) \\
& | & \text{Plus}(\Expr, \Expr) \\
&|& \text{Minus}(\Expr, \Expr) \\
& |& \text{Mult}(\Expr, \Expr) \\
& | & \text{Geq}(\Expr, \Expr) \\
& | & \text{And}(\Expr, \Expr) \\
&|& \text{Or}(\Expr, \Expr) \\
& | & \text{IfThenElse}(\Expr, \Expr, \Expr) \\
&|& \text{Let}(\mathbf{String}, \Expr, \Expr) \\
\end{array}$$


$$\newcommand\translate{\textsf{translateIntoScala}}$$ 
The goal is to implement the function  $\translate(e)$ that inputs a Lettuce Expr $e$ and outputs
a string that is supposed to be a scala program. 

We provide semantics as below:

$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array} (\text{#3})}$$

Note that the output of $\translate$ is a string. 
$$\semRule{}{\translate(\text{Const}(f)) = \textsf{convertToString}(f) }{const}$$

Note that to convert a number to string in scala, use the 'toString' method.

$$\semRule{}{\translate(\text{Ident}(x)) = x }{ident}$$
$$\semRule{}{\translate(\text{ConstTrue}) = true }{true}$$
$$\semRule{}{\translate(\text{ConstFalse}) = false }{false}$$

Note that when you translate subexpressions, you have to make sure to wrap them in curly braces so that
things defined in the scope of one subexpression do not accidently fall into another. <font color="red"> Please follow the conventions below with curly braces. Otherwise you will not pass the tests </font>

$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{Plus}(e1, e2)) = \{ s1 \} + \{ s2 \} }{plus}$$
$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{Minus}(e1, e2)) = \{ s1  \} - \{ s2 \} }{minus}$$
$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{Geq}(e1, e2)) = \{ s1 \} >= \{ s2 \} }{geq}$$
$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{And}(e1, e2)) = \{ s1  \}\ \&\!\&\ \{ s2 \} }{and}$$
$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{Or}(e1, e2)) = \{ s1 \}\ |\!|\ \{ s2 \} }{or}$$
$$\semRule{\translate(e) = s, \translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{IfThenEllse}(e,e1, e2)) = \text{if}\ (\{s\})\ \{ s1 \}\ else\ \{ s2 \} }{ite}$$
$$\semRule{\translate(e1) = s1,\ \translate(e2) = s2}{\translate(\text{Let}(x,e1, e2)) = \text{val}\ x\ =\ \{ s1 \}\ \{ s2 \}  }{let}$$

Whitespaces (space, tab, returns) are ignored by our test cases.

```scala
sealed trait Expr
case class Const(f: Double) extends Expr
case object ConstTrue extends Expr
case object ConstFalse extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Geq(e1: Expr, e2: Expr) extends Expr
case class And(e1: Expr, e2: Expr) extends Expr
case class Or(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x:String, e1: Expr, e2: Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mobject[39m [36mConstTrue[39m
defined [32mobject[39m [36mConstFalse[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
```

<!-- grade_id: cell-1b8c86a80b640a5b, points: 0, grade: false, solution: true -->
```scala
def translateIntoScala(e: Expr): String = e match { 
     case Const(e1) => {
         e1.toString()
     }
    
    case Ident(e1) => {
        e1.toString()
    }
    
    case ConstTrue => {
        "true"
    }
    
    case ConstFalse => {
        "false"
    }
    
    case Plus(e1,e2) =>{
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        s"{${e1_translate}} + {${e2_translate}}"
    }
    
    case Minus(e1,e2) =>{
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        println(s"{${e1_translate}} + {${e2_translate}}")
        s"{${e1_translate}} - {${e2_translate}}"
    }
    
    case Geq(e1,e2) => {
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        println(s"{${e1_translate}} >= {${e2_translate}}")
        s"{${e1_translate}} >= {${e2_translate}}"
    }
    
    case And(e1,e2) => {
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        println(s"{${e1_translate}} | {${e2_translate}}")
        s"{${e1_translate}} && {${e2_translate}}"
    }
    
    case Or(e1,e2) => {
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        println(s"{${e1_translate}} || {${e2_translate}}")
        s"{${e1_translate}} || {${e2_translate}}"
    }
    

    case IfThenElse(e1,e2,e3) => {
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        val e3_translate = {
            translateIntoScala(e3) 
        }
        println(e1_translate)
        println(e2_translate)
        println(e3_translate)
        
        // if ({s}) {s1} else {s2}
        println(s"if (${e1_translate}) {${e2_translate}} else {${e3_translate}}")
        
        
        // if ({s}) {s1} else {s2}
        s"if({${e1_translate}}) {${e2_translate}} else {${e3_translate}}"
    }
    
    case Mult(e1,e2) =>{
        val e1_translate = {
            translateIntoScala(e1) 
        }
        val e2_translate = {
            translateIntoScala(e2)
        }
        println(e1_translate)
        println(e2_translate)
        println(s"{${e1_translate}} * {${e2_translate}}")
        s"{${e1_translate}} * {${e2_translate}}"
    }
    
    case Let(x, e1, e2) => {
        val e2_translate = {
            translateIntoScala(e1)
        }
        val e3_translate = {
            translateIntoScala(e2) 
        }
        
        // val x = {s1} {s2}
        println(">>: ", s"val x = {${e2_translate}} {${e3_translate}}")
        s"val ${x} = {${e2_translate}} {${e3_translate}}"
    }
    
    
    
    
}
```

Output:
```
defined [32mfunction[39m [36mtranslateIntoScala[39m
```

<!-- grade_id: cell-33c6c7e93bdb9232, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Const(10.0)
assert(translateIntoScala(e1) == "10.0")
val e2 = Ident("x")
assert(translateIntoScala(e2) == "x", "Translation of identifier failed")
val e3 = ConstTrue
assert(translateIntoScala(e3) == "true", "Translation of true failed")
val e4 = ConstFalse
assert(translateIntoScala(e4) == "false", "Translation of false failed")
passed(4)
//END TEST
```

Output:
```
Tests Passed (4 points)[36me1[39m: [32mConst[39m = [33mConst[39m([32m10.0[39m)
[36me2[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36me3[39m: [32mConstTrue[39m = ConstTrue
[36me4[39m: [32mConstFalse[39m = ConstFalse
```

```scala
// PLEASE RUN THIS CELL BEFORE TESTING FURTHER
def cleanUpWhiteSpacesInString(st: String): String = 
    st.filterNot( _.isWhitespace )

val tst1 = """
val x = { { 1 } + { 3 } }
     {
        val y = { 2 }
        {
            val z = { 10 }
                { x } + { { y } * { z } }
        }
     }

"""
print(cleanUpWhiteSpacesInString(tst1))

def checkWhitespaceMunged(s1: String, s2: String) : Boolean = {
    cleanUpWhiteSpacesInString(s1) == cleanUpWhiteSpacesInString(s2)
}
```

Output:
```
valx={{1}+{3}}{valy={2}{valz={10}{x}+{{y}*{z}}}}defined [32mfunction[39m [36mcleanUpWhiteSpacesInString[39m
[36mtst1[39m: [32mString[39m = [32m"""
val x = { { 1 } + { 3 } }
     {
        val y = { 2 }
        {
            val z = { 10 }
                { x } + { { y } * { z } }
        }
     }

"""[39m
defined [32mfunction[39m [36mcheckWhitespaceMunged[39m
```

<!-- grade_id: cell-5cda75ffe8ddf0f6, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST

val x = Ident("x")
val y = Ident("y")
val one = Const(1.0)
val two = Const(2.0)

val e1 = Plus(x, y)
val e1Expected = "{ x } + { y }"
assert(checkWhitespaceMunged(translateIntoScala(e1), e1Expected), "Failed test for e1")

val e2 = Plus(x, one)
val e2Expected = "{ x } + { 1.0 }"
assert(checkWhitespaceMunged(translateIntoScala(e2), e2Expected), "Failed test for e2")


val e3 = Minus(x, two)
val e3Expected = "{ x } - { 2.0 }"
assert(checkWhitespaceMunged(translateIntoScala(e3), e3Expected), "Failed test for e3")

val e4 = Plus(x, Minus(y, one))
val e4Expected = "{ x } + { { y } - {1.0} }"
val s4 = translateIntoScala(e4)
println("s4= " + s4)
assert(checkWhitespaceMunged(s4, e4Expected), "Failed test for e4")

val e5 = And( Geq( x, y) , Geq(y, one))
val e5Expected = "{{x} >= {y}} && {{y} >= {1.0}}"
val s5 = translateIntoScala(e5)
println("s5 = " + s5)
assert(checkWhitespaceMunged(s5, e5Expected), "Failed test for e5")


val e6 = Or( Geq( Mult(x, y), Ident("z")) , Geq(y, one))
val e6Expected = "{{ {x} * {y}} >= {z}} || {{y} >= {1.0}}"
val s6 = translateIntoScala(e6)
println("s6 = " + s6)
assert(checkWhitespaceMunged(s6, e6Expected), "Failed test for e6")

passed(6)

//END TEST
```

Output:
```
x
y
x
1.0
x
2.0
{x} + {2.0}
y
1.0
{y} + {1.0}
x
{y} - {1.0}
s4= {x} + {{y} - {1.0}}
x
y
{x} >= {y}
y
1.0
{y} >= {1.0}
{x} >= {y}
{y} >= {1.0}
{{x} >= {y}} | {{y} >= {1.0}}
s5 = {{x} >= {y}} && {{y} >= {1.0}}
x
y
{x} * {y}
{x} * {y}
z
{{x} * {y}} >= {z}
y
1.0
{y} >= {1.0}
{{x} * {y}} >= {z}
{y} >= {1.0}
{{{x} * {y}} >= {z}} || {{y} >= {1.0}}
s6 = {{{x} * {y}} >= {z}} || {{y} >= {1.0}}
Tests Passed (6 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mone[39m: [32mConst[39m = [33mConst[39m([32m1.0[39m)
[36mtwo[39m: [32mConst[39m = [33mConst[39m([32m2.0[39m)
[36me1[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))
[36me1Expected[39m: [32mString[39m = [32m"{ x } + { y }"[39m
[36me2[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m))
[36me2Expected[39m: [32mString[39m = [32m"{ x 
```

<!-- grade_id: cell-46ed1792e2756a80, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val one = Const(1.0)
val two = Const(2.0)
val e5 = And( Geq( x, y) , Geq(y, one))
val e5Expected = "{{x} >= {y}} && {{y} >= {1.0}}"
val iteExpr = IfThenElse(e5, one, y)
val iteExpected = s""" if ({$e5Expected}) {
    1.0
} else {
    y
}
"""
val iteTranslated = translateIntoScala(iteExpr)
println("iteTranslated = " + iteTranslated)
assert(checkWhitespaceMunged(iteTranslated, iteExpected), "Failed test for ite")
passed(5)
//END TEST
```

Output:
```
x
y
{x} >= {y}
y
1.0
{y} >= {1.0}
{x} >= {y}
{y} >= {1.0}
{{x} >= {y}} | {{y} >= {1.0}}
{{x} >= {y}} && {{y} >= {1.0}}
1.0
y
if ({{x} >= {y}} && {{y} >= {1.0}}) {1.0} else {y}
iteTranslated = if({{{x} >= {y}} && {{y} >= {1.0}}}) {1.0} else {y}
Tests Passed (5 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mone[39m: [32mConst[39m = [33mConst[39m([32m1.0[39m)
[36mtwo[39m: [32mConst[39m = [33mConst[39m([32m2.0[39m)
[36me5[39m: [32mAnd[39m = [33mAnd[39m([33mGeq[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mGeq[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m1.0[39m)))
[36me5Expected[39m: [32mString[39m = [32m"{{x} >= {y}} && {{y} >= {1.0}}"[39m
[36miteExpr[39m: [32mIfThenElse[39m = [33mIfThenElse[39m(
  [33mAnd[39m([33mGeq[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mGeq[39m([33mIdent[39m([32m"
```

<!-- grade_id: cell-93ecfa7b29d530af, points: 10, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val one = Const(1.0)
val two = Const(2.0)
val e6 = Or( Geq( Mult(x, y), Ident("z")) , Geq(y, one))
val e6Expected = "{{ {x} * {y}} >= {z}} || {{y} >= {1.0}}"

val letExpr1 = Let("x", Const(10.0) , Plus(x, one))
val letExpr1Expected = """val x = { 10.0 } {
    { x } + { 1.0 } }
"""
val letExprTranslated1 = translateIntoScala(letExpr1)
println("letExprTranslated1 = " + letExprTranslated1)
assert(checkWhitespaceMunged(letExprTranslated1, letExpr1Expected), "Failed test for letExpr1")


val letExpr2 = Let("x", letExpr1 , Let ("y", two, Plus(x, y)))
val letExprTranslated2 = translateIntoScala(letExpr2)
println("letExprTranslated2 = " + letExprTranslated2)
val letExprExpected2 = """ val x = {
                 	   
                val x = {
                 	   10.0
                }
                {
                 	  { x } + {1.0}
                }
            
                }
                {
                 	  
                val y = {
                 	   2.0
                }
                {
                 	  { x } + {y}
                }
            
                }
"""

assert(checkWhitespaceMunged(letExprTranslated2, letExprExpected2), "Failed test for letExpr2")
passed(10)
//END TEST
```

Output:
```
x
1.0
(>>: ,val x = {10.0} {{x} + {1.0}})
letExprTranslated1 = val x = {10.0} {{x} + {1.0}}
x
1.0
(>>: ,val x = {10.0} {{x} + {1.0}})
x
y
(>>: ,val x = {2.0} {{x} + {y}})
(>>: ,val x = {val x = {10.0} {{x} + {1.0}}} {val y = {2.0} {{x} + {y}}})
letExprTranslated2 = val x = {val x = {10.0} {{x} + {1.0}}} {val y = {2.0} {{x} + {y}}}
Tests Passed (10 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mone[39m: [32mConst[39m = [33mConst[39m([32m1.0[39m)
[36mtwo[39m: [32mConst[39m = [33mConst[39m([32m2.0[39m)
[36me6[39m: [32mOr[39m = [33mOr[39m(
  [33mGeq[39m([33mMult[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m)), [33mIdent[39m([32m"z"[39m)),
  [33mGeq[39m([33mIdent[39m([32m"y"[39m), [33mConst[39m([32m1.0[39m))
)
[36me6Expected[39m: [32mString[39m = [32m"{{ {x} * {y}} >= {z}} || {{y} >= {1.0}}"[39m
[36mletExpr1[39m: [32mLet[39m = [33mLet
```

# Problem 2: Let it Go (25 points)

In this problem, we explore the concept of "referential transparency" in lettuce by showing that we can get rid of let bindings.

The main idea is to rewrite Lettuce programs to get rid of Let bindings systematically. 

### Example 1

~~~
let x = 15 in 
   x + x 
~~~

should translate into the expression

~~~
15 + 15
~~~

### Example 2

~~~
let y = 2 * z in 
  let x = 15* y + z in 
    x * x - x
~~~

should become

~~~
 (15 * (2 * z)  + z ) * (15 * (2 * z)  + z ) - ((15 * (2*z) + z))
~~~

Note that we are not evaluating the expression but simply performing substituions in order to
~~~
transform an expression with let bindings into an expression without let bindings that is equivalent.
~~~

For simplicity, we will just work with the following fragment

$$\newcommand\Expr{\mathbf{Expr}}$$
$$ \begin{array}{rcl}
\Expr &\Rightarrow & \text{Const}(\mathbf{Double}) \\
& | & \text{Ident}(\mathbf{String}) \\
& | & \text{Plus}(\Expr, \Expr) \\
& | & \text{IfThenElse}(\Expr, \Expr, \Expr) \\
&|& \text{Let}(\mathbf{String}, \Expr, \Expr) \\
\end{array}$$


$$\newcommand\removeLet{\textsf{removeLetBindings}}$$

The function we define has the following type:
$$\removeLet: \Expr \times \text{Map}[String, \Expr]\ \Rightarrow\ \Expr$$.

In other words, $\removeLet$ takes two arguments, an expression and a _substitution enviromnent_ that maps identifier names to  expressions that are to be substituted for those names, and returns an expression. 

$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array} (\text{#3})}$$

$$\semRule{}{\removeLet(\text{Const}(f), \sigma) = \text{Const}(f) } {const}$$

There are two rules for identifiers. If an identifier exists in the substition environment map, then
carry out the substitution.
$$\semRule{ x \in \mathsf{domain}(\sigma) }{\removeLet(\text{Ident}(x), \sigma) = \sigma(x) } {identifier-substitute}$$

If an identifier does not exist, then leave it unchanged.
$$\semRule{ x \not\in \mathsf{domain}(\sigma) }{\removeLet(\text{Ident}(x), \sigma) = \text{Ident}(x) } {identifier-leave-unchanged}$$

Here are the rules for plus and if then else:

$$\semRule{ \removeLet(e1, \sigma) = e1',\  \removeLet(e2, \sigma) = e2'}{\removeLet(\text{Plus}(e1, e2), \sigma) = \text{Plus}(e1', e2') }{plus}$$

$$\semRule{ \removeLet(e, \sigma) = e',\ \removeLet(e1, \sigma) = e1',\ \removeLet(e2, \sigma) = e2'}{\removeLet(\text{IfThenElse}(e, e1, e2), \sigma) = \text{IfThenElse}(e', e1', e2') }{ite}$$


Finally, the rule for let bindings is this:

$$\semRule{ \removeLet(e1, \sigma_{old} ) = e1',\ \color{red}{\sigma_{new}} = \sigma_{old} \circ [ x \mapsto e1' ] }{ \removeLet(\text{Let}(e, e1, e2), \sigma_{old}) = \removeLet(e2, \color{red}{\sigma_{new}}) }{let-bindings}$$

   
Please make sure you run this cell below: <font color="red"> Warning expressions Expr is being redefined here. If you want to work on P1, make sure that the definitions for P1 are run from the appropriate cell for that problem. Othewrwise you will get strange errors. </font>

```scala
sealed trait Expr
case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x:String, e1: Expr, e2: Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
```

<!-- grade_id: cell-aab368a23cbafdcd, points: 0, grade: false, solution: true -->
```scala
def removeLet(e: Expr, env: Map[String, Expr]): Expr = e match{
    case Const(e1) => {
        Const(e1)
    }
    
    case Ident(e1) => {
        if (env contains e1){
            env(e1)
        }else{
            Ident(e1)
        }
    }
    
    case Plus(e1,e2) => {
        val v1 = removeLet(e1, env)
        val v2 = removeLet(e2, env) 
        Plus(v1, v2) 
    }
    
    case IfThenElse(e1, e2, e3) => {
        val v1 = removeLet(e1, env)
        val v2 = removeLet(e2, env) 
        val v3 = removeLet(e3, env) 
        IfThenElse(v1,v2,v3)
    }
    
    case Let(x, e1, e2) => {
        val v1 = removeLet(e1, env) 
        // update the map 
        val newEnv = env.updated(x,v1);
        removeLet(e2,newEnv)
    }
    case _ => {
        throw new IllegalArgumentException("something went wrong")
       
    }
}
```

Output:
```
defined [32mfunction[39m [36mremoveLet[39m
```

$$\semRule{ \removeLet(e1, \sigma_{old} ) = e1',\ \color{red}{\sigma_{new}} = \sigma_{old} \circ [ x \mapsto e1' ] }{ \removeLet(\text{Let}(e, e1, e2), \sigma_{old}) = \removeLet(e2, \color{red}{\sigma_{new}}) }{let-bindings}$$


<!-- grade_id: cell-887aa26e1ddd9805, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val one = Const(1.0)

val emptyEnv: Map[String, Expr] = Map.empty
assert(removeLet(Const(1.0), emptyEnv) == Const(1.0), "test failed on input Const(1.0)")

val e1 = Plus(x, one)
val e1res = removeLet(e1, emptyEnv)
println("e1res = "  + e1res)
assert( e1res == e1, "test e1 failed")

val e2 = IfThenElse(Plus(x, one), Plus(x, one), Plus(y, x))
val e2res = removeLet(e2, emptyEnv)
println("e2res = "  + e2res)
assert( e2res == e2, "test e2 failed")

passed(5)

//END TEST
```

Output:
```
e1res = Plus(Ident(x),Const(1.0))
e2res = IfThenElse(Plus(Ident(x),Const(1.0)),Plus(Ident(x),Const(1.0)),Plus(Ident(y),Ident(x)))
Tests Passed (5 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mone[39m: [32mConst[39m = [33mConst[39m([32m1.0[39m)
[36memptyEnv[39m: [32mMap[39m[[32mString[39m, [32mExpr[39m] = [33mMap[39m()
[36me1[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m))
[36me1res[39m: [32mExpr[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m))
[36me2[39m: [32mIfThenElse[39m = [33mIfThenElse[39m(
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m)),
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m)),
  [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"x"[39m))
)
[36me2res[39m: [32mExpr[39m = [33mIfThe
```

<!-- grade_id: cell-46ca6228a8df2cd8, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val emptyEnv: Map[String, Expr] = Map.empty
val env1 = emptyEnv + ("x" -> Const(5.0))
assert(removeLet(Ident("x"), emptyEnv) == Ident("x"), "test failed on input Ident(x), empty env")
assert(removeLet(Ident("x"), env1) == Const(5.0), "test failed on input Ident(x), env that maps x to Const(5.0)")
//END TEST
passed(5)

```

Output:
```
Tests Passed (5 points)[36memptyEnv[39m: [32mMap[39m[[32mString[39m, [32mExpr[39m] = [33mMap[39m()
[36menv1[39m: [32mMap[39m[[32mString[39m, [32mExpr[39m] = [33mMap[39m([32m"x"[39m -> [33mConst[39m([32m5.0[39m))
```

<!-- grade_id: cell-a5ff8d7a2503a9c7, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val z = Ident("z")
val e1 = Plus(x, Const(1.0))
val e2 = Plus(x, y)
val e3 = Plus(x, Plus(x, Plus(y, z)))

val emptyEnv: Map[String, Expr] = Map.empty


val letExpr1 = Let ("x", Const(10.0), e1)
val letExprResult1 = removeLet(letExpr1, emptyEnv)
print("letExpr1 result = " + letExprResult1)
assert(letExprResult1 == Plus(Const(10.0), Const(1.0)), "test 1 failed for let")

passed(4)

//END TEST
```

Output:
```
letExpr1 result = Plus(Const(10.0),Const(1.0))Tests Passed (4 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mz[39m: [32mIdent[39m = [33mIdent[39m([32m"z"[39m)
[36me1[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m))
[36me2[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))
[36me3[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"z"[39m))))
[36memptyEnv[39m: [32mMap[39m[[32mString[39m, [32mExpr[39m] = [33mMap[39m()
[36mletExpr1[39m: [32mLet[39m = [33mLet[39m([32m"x"[39m, [33mConst[39m([32m10.0[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m)))
[36mletExprResult1[39m: [32mExpr[39m = [33mPlus[3
```

<!-- grade_id: cell-53d865eb138a1cae, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val z = Ident("z")
val e1 = Plus(x, Const(1.0))
val e2 = Plus(x, y)
val e3 = Plus(x, Plus(x, Plus(y, z)))

val emptyEnv: Map[String, Expr] = Map.empty


val letExpr2 = Let ("y", e1, e2)
val letExprResult2 = removeLet(letExpr2, emptyEnv)
print("letExpr1 result = " + letExprResult1)
assert(letExprResult2 == Plus(Ident("x"), Plus(Ident("x"), Const(1.0))), "test 1 failed for let")

passed(4)
//END TEST
```

Output:
```
letExpr1 result = Plus(Const(10.0),Const(1.0))Tests Passed (4 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mz[39m: [32mIdent[39m = [33mIdent[39m([32m"z"[39m)
[36me1[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m))
[36me2[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))
[36me3[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"z"[39m))))
[36memptyEnv[39m: [32mMap[39m[[32mString[39m, [32mExpr[39m] = [33mMap[39m()
[36mletExpr2[39m: [32mLet[39m = [33mLet[39m(
  [32m"y"[39m,
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mConst[39m([32m1.0[39m)),
  [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"y"[39m))
)

```

<!-- grade_id: cell-f2d1452831f4d53a, points: 7, grade: true, solution: false -->
```scala
//BEGIN TEST
val x = Ident("x")
val y = Ident("y")
val z = Ident("z")
val e1 = Plus(Plus(x, x), Plus(x, x))
val e2 = Plus(Plus(x, x), Plus(y, y))
val e3 = Plus(x, Plus(x, Plus(y, z)))

val emptyEnv: Map[String, Expr] = Map.empty

val letExpr3 = Let("x", Const(10.0), Let("y", e1, Let("z", e2, e3)))
val letExpr3Res = removeLet(letExpr3, emptyEnv)
println("let expr3 result = " + letExpr3Res)
assert(letExpr3Res ==  Plus( Const(10.0),  Plus(  Const(10.0), Plus( Plus(Plus(Const(10.0), Const(10.0)), Plus(Const(10.0), Const(10.0))), Plus(
        Plus(Const(10.0), Const(10.0)), Plus( Plus(Plus(Const(10.0), Const(10.0)), Plus(Const(10.0), Const(10.0))),Plus(Plus(Const(10.0), Const(10.0)), Plus(Const(10.0), Const(10.0))))
      )))), "test for let expr 3 failed")

passed(7)
//END TEST
```

Output:
```
let expr3 result = Plus(Const(10.0),Plus(Const(10.0),Plus(Plus(Plus(Const(10.0),Const(10.0)),Plus(Const(10.0),Const(10.0))),Plus(Plus(Const(10.0),Const(10.0)),Plus(Plus(Plus(Const(10.0),Const(10.0)),Plus(Const(10.0),Const(10.0))),Plus(Plus(Const(10.0),Const(10.0)),Plus(Const(10.0),Const(10.0))))))))
Tests Passed (7 points)[36mx[39m: [32mIdent[39m = [33mIdent[39m([32m"x"[39m)
[36my[39m: [32mIdent[39m = [33mIdent[39m([32m"y"[39m)
[36mz[39m: [32mIdent[39m = [33mIdent[39m([32m"z"[39m)
[36me1[39m: [32mPlus[39m = [33mPlus[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)), [33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)))
[36me2[39m: [32mPlus[39m = [33mPlus[39m([33mPlus[39m([33mIdent[39m([32m"x"[39m), [33mIdent[39m([32m"x"[39m)), [33mPlus[39m([33mIdent[39m([32m"y"[39m), [33mIdent[39m([32m"y"[39m)))
[36me3[39m: [32mPlus[39m = [33mPlus[39m([33mIdent[39m([32m"x"[39m), 
```

```scala

```