```scala

```

# CSCI 3155 Assignment 1 : Scala Programming

This assignment asks you to write scala programs. Fill in the programs below. Use the test cases provided to test them. You are also encouraged to write your own test cases.

**YOUR NAME HERE**

```scala
// TEST HELPER: Do not EDIT this cell please.
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage(cond: Boolean, testID: String) = { 
    assert(cond, s"Test $testID Failed"); 
    println(s"Test $testID Passed!")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 1 (5 points)

Write a scala function titled `isPalindrome (s)` that tests whether a given input string `s` is a palindrome.

The following string operations will be helpful.

1. `s.size` gives you the size of the string.
2. `s(i)` gives you the `ith` character for `i = 0, ... , s.size-1`
3. For loops in scala are described here: https://www.tutorialspoint.com/scala/scala_for_loop.htm
   Pay careful attention between the to and until forms of the for loop.

<!-- grade_id: cell-794743da17722228, points: 0, grade: false, solution: true -->
```scala
def isPalindrome(s: String): Boolean = { 
    if (s == "" || s.length == 1){
        return true
    }
    else if (s.length == 2){
        if (s(0) == s(1)){
            return true
        }
        else{false}
    }
    else{
        var backwardsString = "" 
        for (itr <- s.length-1 to 0 by -1){
            
            backwardsString = backwardsString + s(itr)
        }
        println(backwardsString)
        println(s)
        if (backwardsString == s){
            true
        }else{
            false
        }
    }
 }
```

Output:
```
defined [32mfunction[39m [36misPalindrome[39m
```

<!-- grade_id: cell-082571ccc4ca97d0, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
assert(isPalindrome("hellolleh"), "Test 1 Failed")
assert(!isPalindrome("bababablacksheep"), "Test 2 Failed")
assert(isPalindrome(""), "Test 3 Failed")
assert(isPalindrome("a"), "Test 4 Failed")
assert(isPalindrome("xxxxxxxx"), "Test 5 Failed")
assert(isPalindrome("ablewasiereisawelba"), "Test 6 Failed")
passed(5)
//END TEST
```

Output:
```
hellolleh
hellolleh
peehskcalbababab
bababablacksheep
xxxxxxxx
xxxxxxxx
ablewasiereisawelba
ablewasiereisawelba

*** Tests Passed (5 points) ***
```

## Problem 2 (5 points)

Implement a `isPalindromeRec` function that checks if a given string is a palindrome, but without using mutable `var`, or without using `for`/`while` loops.

**Hint:** You should consult Scala documentation on strings to figure out:
1. How do we get first and last characters of a string.
2. How do we get substring between first and last character.
3. How to find out if a string is empty or a single character.
https://www.tutorialspoint.com/scala/scala_strings.htm

<!-- grade_id: cell-80da0f71ccafab19, points: 0, grade: false, solution: true -->
```scala
def isPalindromeRec(s: String): Boolean = { 
    if (s == "" || s.length == 1){
        return true
    }
    else if (s.length == 2){
        if (s(0) == s(1)){
            return true
        }
        else{
            return false
        }
    }
    else{
        isPalindrome(s.substring(1,s.length-1))
    }
 }
```

Output:
```
defined [32mfunction[39m [36misPalindromeRec[39m
```

<!-- grade_id: cell-ef781cc6d55287ba, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST

testWithMessage(isPalindromeRec("hellolleh") == true, "1")
testWithMessage(isPalindromeRec("bababablacksheep") == false, "2")
testWithMessage(isPalindromeRec("") == true, "3")
testWithMessage(isPalindromeRec("a") == true, "4")
testWithMessage(isPalindromeRec("xxxxxxxx") == true, "5")
testWithMessage(isPalindromeRec("ablewasiereisawelba") == true, "6")
passed(5)
//END TEST
```

Output:
```
ellolle
ellolle
Test 1 Passed!
eehskcalbababa
ababablackshee
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
xxxxxx
xxxxxx
Test 5 Passed!
blewasiereisawelb
blewasiereisawelb
Test 6 Passed!

*** Tests Passed (5 points) ***
```

## Problem 3 (5 points)

A function $f: \mathbb{Z} \rightarrow \mathbb{Z}$ over integers is a odd function if for all numbers $x \in \mathbb{Z}$, $f(-x) = -f(x)$.

As an example $f(x) = 5x$ is an odd function. However, $f(x) = x^2$ is not.


Write a scala function `isOddFunction(f, n)` that checks if $f(-z) = -f(z)$ for every number $z$ within the range $1, \ldots, n$. The function should return `true` if $f$ is indeed an odd function and `false` otherwise. You may use a for loop for this problem. 

<!-- grade_id: cell-096de6a7e20071c9, points: 0, grade: false, solution: true -->
```scala
def isOddFunction(f: Int => Int, n: Int): Boolean = {
    for ( itr <- 1 to n ){
        if ((-f(itr)) != f(-itr)){
            return false
        }
    }
    return true
}
```

Output:
```
defined [32mfunction[39m [36misOddFunction[39m
```

<!-- grade_id: cell-af1b5ad4b16aa37c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
/* Define functions for the tests */
val f1: Int => Int = (x) => (x+10)
val f2: Int => Int = (x) => (-2*x)
def f3(x: Int): Int = x*x
val f4: Int => Int = (x) => {
    if (x < 10)  - 2 * x
    else 2 * x
}

/* Some test cases */
testWithMessage(!isOddFunction(f1, 10), "1")
testWithMessage(isOddFunction(f2, 100), "2")
testWithMessage(!isOddFunction(f3, 100), "3")
testWithMessage(isOddFunction(f4, 8), "4")
testWithMessage(!isOddFunction(f4, 15), "5")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
[36mf1[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd6$Helper$$Lambda$2195/2147070144@ff8e300
[36mf2[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd6$Helper$$Lambda$2196/1347833030@2e8d8b60
defined [32mfunction[39m [36mf3[39m
[36mf4[39m: [32mInt[39m => [32mInt[39m = ammonite.$sess.cmd6$Helper$$Lambda$2197/1050814025@2d646640
```

## Problem 4 (10 points)

We would like you to implement the methods in the class Complex  that is supposed to represent complex numbers $a + i b$ wherein $a, b$ are double precision and $i$ is the square root of unity.

Implement the missing methods below.

You should read this blog post about how to convert between numeric types.

https://alvinalexander.com/scala/how-to-convert-between-numeric-types-in-scala-int-long-float-double


<!-- grade_id: cell-a410d01fa666b077, points: 0, grade: false, solution: true -->
```scala
class Complex(val real: Double, val im:Double) {
    // Implement a pretty printer
    override def toString(): String = {
        s"$real + i $im"
    }
    
    
    /* The modulus of a complex number is given by sqrt(real*real + im * im) */
    def modulus(): Double = {
        return scala.math.sqrt((this.real * this.real) + (this.im * this.im))
    }
    
    
    /* The sum of two complex numbers, we are overriding the + operator here.
       How to add complex numbers? (a+ib) + (c+id)  = (a+c) + i (b+d)
    */
    def + (r: Complex) : Complex = {
        val newComplex = new Complex((this.real + r.real),(this.im + r.im))
        newComplex
    }
    
    /* The difference between two  complex numbers.
       How to subtract complex numbers? (a+ib) - (c+id)  = (a-c) + i (b-d)
    */
    def - (r: Complex): Complex = {
        val newComplex = new Complex((this.real - r.real),(this.im - r.im))
        newComplex
    }
    
    /* Product of complex numbers. 
       How to multiply complex numbers? (a+ ib)* (c+id) = (ac - bd) + i (ad + bc)
    */
    def * (r: Complex): Complex = {
        ??? // YOUR CODE HERE
    }
    
    /* Given a complex number (a + ib) its inverse is defined as
      1/(a^2 + b^2) * ( a - ib), 
       provided a != 0 or b != 0. 
       Do not worry about error handling for now. 
    */
    def inverse() : Complex = {
        ??? // YOUR CODE HERE
    }
    
    /* Compare two complex numbers for "approximate" equality within epsilon.
        I.e, (a+ib) approximately equals (c+id) iff  
                         -epsilon <= (a -c) <= epsilon and
                         -epsilon <= (b-d) <= epsilon.
        As a rule, you should avoiding equality comparisons of floating point numbers: 
        find out why! 
    */
    
    def complexApproxEquals(r: Complex, epsilon: Double = 0.01): Boolean = {
        (r.real - epsilon <= real && real  <= r.real + epsilon) && 
        (r.im - epsilon <= im && im <= r.im + epsilon)
    }
    
    // Ignore this method for now -- it will be clearer soon
    override def equals(r: Any): Boolean = r match {
        case r: Complex => this.complexApproxEquals(r, 0.01) 
        case _ => false
    }
    // END Ignore   
}
```

Output:
```
defined [32mclass[39m [36mComplex[39m
```

<!-- grade_id: cell-1e343fa83dc6232e, points: 3, grade: true, solution: false -->
```scala
//BEGIN TEST
/* Test modulus */

val c1 = new Complex(1,1)
val c1mod = c1.modulus()
testWithMessage(math.sqrt(1.999) <= c1mod && c1mod <= math.sqrt(2.001), "modulus-test-1")
val c2 = new Complex(3, 4)
val c2mod = c2.modulus()
testWithMessage(4.99 <= c2mod && c2mod  <= 5.01, "modulus-test-2" )
passed(3)
//END TEST
```

Output:
```
Test modulus-test-1 Passed!
Test modulus-test-2 Passed!

*** Tests Passed (3 points) ***
[36mc1[39m: [32mComplex[39m = 1.0 + i 1.0
[36mc1mod[39m: [32mDouble[39m = [32m1.4142135623730951[39m
[36mc2[39m: [32mComplex[39m = 3.0 + i 4.0
[36mc2mod[39m: [32mDouble[39m = [32m5.0[39m
```

<!-- grade_id: cell-405e5cef64a02f3a, points: 3, grade: true, solution: false -->
```scala
//BEGIN TEST
/* Test Addition and Subtraction */
val c1 = new Complex(1,1)
val c2 = new Complex(1,2)
val c3 = new Complex(2,3)

val c4 = c1 + c2
val c5 = c2 - c3
val c6 = c1 + c3
val c7 = c1 - c2


def approxEq(c: Complex, r1: Double, r2: Double) = {
    c.real <= r1 + 0.001 && c.real >= r1 - 0.001 &&
    c.im >= r2 - 0.001 && c.im <= r2 +0.001
}
testWithMessage(approxEq(c4, 2.0, 3.0), "addition-1")
testWithMessage(approxEq(c5, -1.0, -1.0), "subtraction-1")
testWithMessage(approxEq(c6, 3.0, 4.0), "addition-2")
testWithMessage(approxEq(c7, 0.0, -1.0), "subtraction-2")
passed(3)
//END TEST
```

Output:
```
Test addition-1 Passed!
Test subtraction-1 Passed!
Test addition-2 Passed!
Test subtraction-2 Passed!

*** Tests Passed (3 points) ***
[36mc1[39m: [32mComplex[39m = 1.0 + i 1.0
[36mc2[39m: [32mComplex[39m = 1.0 + i 2.0
[36mc3[39m: [32mComplex[39m = 2.0 + i 3.0
[36mc4[39m: [32mComplex[39m = 2.0 + i 3.0
[36mc5[39m: [32mComplex[39m = -1.0 + i -1.0
[36mc6[39m: [32mComplex[39m = 3.0 + i 4.0
[36mc7[39m: [32mComplex[39m = 0.0 + i -1.0
defined [32mfunction[39m [36mapproxEq[39m
```

<!-- grade_id: cell-d8bb7ad100847752, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
val c1 = new Complex(1,1)
val c2 = new Complex(1,2)
val c3 = new Complex(2,3)

val c8 = c1 * c2
val c9 = c1 * c3
val c10 = c1 - c2 * c3 + c2*c2

def approxEq(c: Complex, r1: Double, r2: Double) = {
    c.real <= r1 + 0.001 && c.real >= r1 - 0.001 &&
    c.im >= r2 - 0.001 && c.im <= r2 +0.001
}
testWithMessage(approxEq(c8, -1.0, 3.0), "multiplication-1")
testWithMessage(approxEq(c9, -1.0, 5.0), "multiplication-2")
testWithMessage(approxEq(c10, 2.0, -2.0), "arithmetic-1")
passed(2)
//END TEST
```

Output:
```
scala.NotImplementedError: an implementation is missing
  scala.Predef$.$qmark$qmark$qmark(Predef.scala:288)
  ammonite.$sess.cmd7$Helper$Complex.$times(cmd7.sc:34)
  ammonite.$sess.cmd10$Helper.<init>(cmd10.sc:5)
  ammonite.$sess.cmd10$.<init>(cmd10.sc:7)
  ammonite.$sess.cmd10$.<clinit>(cmd10.sc:-1)
```

<!-- grade_id: cell-4c586ade24b753a2, points: 2, grade: true, solution: false -->
```scala
//BEGIN TEST
// Test inverse
val c1 = new Complex(1,1)
val c2 = new Complex(1,2)
val c3 = new Complex(2,3)
val c9 = c1.inverse()
val c10 = c2.inverse()
val c11 = c3.inverse()

val c9_expected = new Complex(0.5,-0.5)
val c10_expected = new Complex(0.2,-0.4)
val c11_expected = new Complex(0.15385, -0.23077)

testWithMessage(c9 == c9_expected , "c1-inverse-test")
testWithMessage(c10 == c10_expected , "c2-inverse-test")
testWithMessage(c11 == c11_expected , "c3-inverse-test")
passed(2)
//END TEST
```

## Problem 5 (20 points)

Newton invented the Newton-Raphson method for solving an equation. We are going to ask you to write some code to solve equations.

To find roots of an equation of the form

$$ x^2 - 3x + 2 = 0$$

we start from an initial guess at the solution: say $x_0 = 4.5$
Each time we have the $i^{th}$ guess $x_i$, we update it as

$$ x_{i+1} = x_i - \frac{f(x_i)}{f'(x_i)}, $$

where $f'(x) = \frac{df(x)}{dx}$.

For our equation, $f(x) = x^2 - 3x +2$ and $f'(x) = 2 x - 3$. Thus, our update equation is 
$$ x_{i+1} = x_i - \frac{x_i^2 - 3 x_i + 2}{2 x_i - 3}$$.

We stop whenever $|f(x_i)| \leq 10^{-8}$: i.e, we are very close to a root of the function. Gory details are here for the curious ones: http://www.math.ubc.ca/~anstee/math104/newtonmethod.pdf

__(A, 5 points)__ Write a scala function `calculateRoot(x0: Double)` that takes in the initial guess $x_0$ as input and calculates the root for the polynomial $x^2 - 3 x + 2$. Your code should use a While loop.

Caution: To calculate $x^2$, you should use `math.pow(x,2)` or `x*x` in Scala. `x^2` does not work in scala since `^` operator actually does XOR.

<!-- grade_id: cell-2dd21341b51d86ad, points: 0, grade: false, solution: true -->
```scala
def calculateRoot(x0: Double): Double = { 
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    def df(x: Double) = { 2* x - 3}
    // Your code here: use f(x) and df(x) provided already.
    var solution = x0
    while ((f(solution)).abs > scala.math.pow(10,-8)){
        solution = solution - ( f(solution) / df(solution))
    }
    return solution
}
```

Output:
```
defined [32mfunction[39m [36mcalculateRoot[39m
```

<!-- grade_id: cell-2e7493c157d154dc, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
def isRootCorrect(x: Double): Boolean = {
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    val y = calculateRoot(x)
    math.abs(f(y)) <= 1e-08
}

testWithMessage(isRootCorrect(3.5), "1")
testWithMessage(isRootCorrect(4.6), "2")
testWithMessage(isRootCorrect(-3), "3")
testWithMessage(isRootCorrect(-3.5), "4")
testWithMessage(isRootCorrect(1.3), "5")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
defined [32mfunction[39m [36misRootCorrect[39m
```

__(B, 5 points)__ Implement the function `calculateRootRec` that uses a recursion instead of a while loop to compute the root. You should __not__ use a mutable `var` in your code. 

<!-- grade_id: cell-57d2ad2f5ba44329, points: 0, grade: false, solution: true -->
```scala
def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
def df(x: Double) = { 2* x - 3}


def calculateRootRec(x: Double): Double = {
    if ((f(x)).abs <= scala.math.pow(10,-8)){
        x
    }else{
        val nextElement = x - (f(x) / df(x))
        calculateRootRec(nextElement)
    }
}
```

Output:
```
defined [32mfunction[39m [36mf[39m
defined [32mfunction[39m [36mdf[39m
defined [32mfunction[39m [36mcalculateRootRec[39m
```

<!-- grade_id: cell-49bff1b838e3e910, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
def isRootCorrectRec(x: Double): Boolean = {
    val y = calculateRootRec(x)
    (math.abs(y - 2.0) <= 1e-06) || (math.abs(y - 1.0) <= 1e-06)
}

testWithMessage(isRootCorrectRec(3.5), "1")
testWithMessage(isRootCorrectRec(4.6), "2")
testWithMessage(isRootCorrectRec(-3), "3")
testWithMessage(isRootCorrectRec(-3.5), "4")
testWithMessage(isRootCorrectRec(1.3), "5")
passed(5)
//END TEST
```

Output:
```
Test 1 Passed!
Test 2 Passed!
Test 3 Passed!
Test 4 Passed!
Test 5 Passed!

*** Tests Passed (5 points) ***
defined [32mfunction[39m [36misRootCorrectRec[39m
```

__(C, 10 points)__ Write a generalized Newton-Raphson function `solveEquationNewtonRaphson` to solve any 
equation $f(x) = 0$ wherein the functions $f$ and $f'(x)$ are given as inputs. Do not use while loops or mutable `var` in your code.



<!-- grade_id: cell-d053a7e0f075bd32, points: 0, grade: false, solution: true -->
```scala
def solveEquationNewtonRaphson(x: Double, f: Double => Double, df: Double => Double): Double = {
   ??? // YOUR CODE HERE
}
```

<!-- grade_id: cell-4a86f811b0471af9, points: 2, grade: true, solution: false -->
```scala
// Test Cases
def f1(x: Double) = 2.0 * math.pow(x,3) - math.pow(x, 2) - 1.0
def df1(x: Double) = 6.0 * math.pow(x, 2) - 2 * x 
val x1 = solveEquationNewtonRaphson(3.0, f1, df1)
testWithMessage(math.abs(f1(x1)) <= 1e-06, "1")
passed(2)
```

<!-- grade_id: cell-dd5de96626a8f24c, points: 2, grade: true, solution: false -->
```scala
def f2(x: Double) = math.sin(x) - 0.5 * x 
def df2(x: Double) = math.cos(x) - 0.5
val x2 = solveEquationNewtonRaphson(0.7, f2, df2)
testWithMessage(math.abs(f2(x2)) <= 1e-06, "2")
passed(2)
```

<!-- grade_id: cell-24787c8854b20def, points: 2, grade: true, solution: false -->
```scala
def f3(x: Double) = math.log(x) - 2 
def df3(x: Double) = 1/x 
val x3 = solveEquationNewtonRaphson(2.0, f3, df3)
testWithMessage(math.abs(f3(x3)) <= 1e-06, "3")
passed(2)
```

<!-- grade_id: cell-08b0bc89a3d1d886, points: 2, grade: true, solution: false -->
```scala
def f4(x: Double) = math.atan(x) - 1
def df4(x: Double) = 1.0/(1.0 + x*x)
val x4 = solveEquationNewtonRaphson(3.0, f4, df4)
testWithMessage(math.abs(f4(x4)) <= 1e-06, "4")
passed(2)
```

<!-- grade_id: cell-e858482df5178f7d, points: 2, grade: true, solution: false -->
```scala
def f5(x: Double) = math.cos(x)* math.cos(x) - 1
def df5(x: Double) = -2.0 * math.sin(x) * math.cos(x)
val x5 = solveEquationNewtonRaphson(3.5, f5, df5)
testWithMessage(math.abs(f5(x5)) <= 1e-06, "5")
passed(2)
```

-- That is all for now, folks! -- 